# Pandas NBgrader exercises

In this notebook you will find the 6 deliverable exercises regarding pandas. This notebook is autograded so please dont modify any cell except to add your solution. Modifying this notebook may result in you not getting a grade.

```python
import pandas as pd
import numpy as np
```

```python
data = {'animal': ['cat', 'cat', 'snake', 'dog', 'dog', 'cat', 'snake', 'cat', 'dog', 'dog'],
        'age': [2.5, 3, 0.5, np.nan, 5, 2, 4.5, np.nan, 7, 3],
        'visits': [1, 3, 2, 3, 2, 3, 1, 1, 2, 1],
        'priority': ['yes', 'yes', 'no', 'yes', 'no', 'no', 'no', 'yes', 'no', 'no']}

labels = ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j']


df = pd.DataFrame(data, index=labels)

```

**13.** Change the age in row 'f' to 1.5.

<!-- grade_id: cell-f8d4a64dd0fb6112, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df.loc['f', 'age'] = 1.5
df
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-6-bbcf82e8217f> in <module>
      2 df.loc['f', 'age'] = 1.5
      3 df
----> 4 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-675b0274a23f22c9, points: 1, grade: true, solution: false -->
```python
assert df['age']['f'] == 1.5
```

**18.** Sort `df` first by the values in the 'age' in *decending* order, then by the value in the 'visit' column in *ascending* order.

<!-- grade_id: cell-bb44e438722da485, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df.sort_values(by=['age', 'visits'], ascending=[False, True])
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-8-b7f966c595f2> in <module>
      1 # YOUR CODE HERE
      2 df.sort_values(by=['age', 'visits'], ascending=[False, True])
----> 3 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-9fa22615a21cf3f1, points: 1, grade: true, solution: false -->
```python
b = True

for i in range(df.shape[0]-1):
    if df['age'].values[i] < df['age'].values[i]:
        b = False
        break
    
assert b
    
```

**23.** Given a DataFrame of numeric values, say
```python
df = pd.DataFrame(np.random.random(size=(5, 3))) # a 5x3 frame of float values
```

how do you subtract the row mean from each element in the row?

<!-- grade_id: cell-5137285459670e80, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df.sub(df.mean(axis=1), axis=0)
raise NotImplementedError()
```

Output:
```
~\anaconda3\lib\site-packages\pandas\core\ops\array_ops.py in _na_arithmetic_op(left, right, op, is_cmp)
    147             #  will handle complex numbers incorrectly, see GH#32047
    148             raise
--> 149         result = _masked_arith_op(left, right, op)
    150 
    151     if is_cmp and (is_scalar(result) or result is NotImplemented):

~\anaconda3\lib\site-packages\pandas\core\ops\array_ops.py in _masked_arith_op(x, y, op)
     89         if mask.any():
     90             with np.errstate(all="ignore"):
---> 91                 result[mask] = op(xrav[mask], yrav[mask])
     92 
     93     else:

TypeError: unsupported operand type(s) for -: 'str' and 'float'
```

<!-- grade_id: cell-8ceb24481d7d7a28, points: 2, grade: true, solution: false -->
```python
assert df.shape == (5,3)
assert df[0].max() <= 1
assert df[0].min() >= -1
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-12-92366bd46582> in <module>
----> 1 assert df.shape == (5,3)
      2 assert df[0].max() <= 1
      3 assert df[0].min() >= -1

AssertionError:
```

**24.** Suppose you have DataFrame with 10 columns of real numbers, for example:

```python
df = pd.DataFrame(np.random.random(size=(5, 10)), columns=list('abcdefghij'))
```
Which column of numbers has the smallest sum? (Find that column's label.)

<!-- grade_id: cell-7830a1e1958310b9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df.sum().idxmin()
raise NotImplementedError()
```

Output:
```
~\anaconda3\lib\site-packages\pandas\core\series.py in idxmin(self, axis, skipna, *args, **kwargs)
   2101         """
   2102         skipna = nv.validate_argmin_with_skipna(skipna, args, kwargs)
-> 2103         i = nanops.nanargmin(self._values, skipna=skipna)
   2104         if i == -1:
   2105             return np.nan

~\anaconda3\lib\site-packages\pandas\core\nanops.py in _f(*args, **kwargs)
     66             if any(self.check(obj) for obj in obj_iter):
     67                 f_name = f.__name__.replace("nan", "")
---> 68                 raise TypeError(
     69                     f"reduction operation '{f_name}' not allowed for this dtype"
     70                 )

TypeError: reduction operation 'argmin' not allowed for this dtype
```

<!-- grade_id: cell-a0daad4562804d2a, points: 2, grade: true, solution: false -->
```python
assert idx in 'abcdefghij'
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-14-58cd645e81a2> in <module>
----> 1 assert idx in 'abcdefghij'

NameError: name 'idx' is not defined
```

The next three puzzles are slightly harder...

**26.** You have a DataFrame that consists of 10 columns of floating--point numbers. Suppose that exactly 5 entries in each row are NaN values. For each row of the DataFrame, find the *column* which contains the *third* NaN value.

(You should return a Series of column labels.)

<!-- grade_id: cell-e4f640a69590b406, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
(df.isnull().cumsum(axis=1) == 3).idxmax(axis=1)
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-15-89a74837bf76> in <module>
      1 # YOUR CODE HERE
      2 (df.isnull().cumsum(axis=1) == 3).idxmax(axis=1)
----> 3 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-182477a2beb50458, points: 3, grade: true, solution: false -->
```python
assert indexes == [3, 8, 5, 6, 8]
assert res.shape == (5,)
assert list(res.values) == [3, 8, 5, 6, 8]
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-16-7419831cb40b> in <module>
----> 1 assert indexes == [3, 8, 5, 6, 8]
      2 assert res.shape == (5,)
      3 assert list(res.values) == [3, 8, 5, 6, 8]

NameError: name 'indexes' is not defined
```

**27.** A DataFrame has a column of groups 'grps' and and column of numbers 'vals'. For example: 

```python
df = pd.DataFrame({'grps': list('aaabbcaabcccbbc'), 
                   'vals': [12,345,3,1,45,14,4,52,54,23,235,21,57,3,87]})
```
For each *group*, find the sum of the three greatest values.

<!-- grade_id: cell-da30c3e27df05854, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
df.groupby('grps')['vals'].nlargest(3).sum(level=0)
raise NotImplementedError()
```

Output:
```
~\anaconda3\lib\site-packages\pandas\core\groupby\groupby.py in __init__(self, obj, keys, axis, level, grouper, exclusions, selection, as_index, sort, group_keys, squeeze, observed, mutated, dropna)
    558             from pandas.core.groupby.grouper import get_grouper
    559 
--> 560             grouper, exclusions, obj = get_grouper(
    561                 obj,
    562                 keys,

~\anaconda3\lib\site-packages\pandas\core\groupby\grouper.py in get_grouper(obj, key, axis, level, sort, observed, mutated, validate, dropna)
    809                 in_axis, name, level, gpr = False, None, gpr, None
    810             else:
--> 811                 raise KeyError(gpr)
    812         elif isinstance(gpr, Grouper) and gpr.key is not None:
    813             # Add key to exclusions

KeyError: 'grps'
```

<!-- grade_id: cell-dd5c34897ebaa573, points: 3, grade: true, solution: false -->
```python
assert res.values[0] == 409
assert res.values[1] == 156
assert res.values[2] == 345
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-18-c3c67b07c1be> in <module>
----> 1 assert res.values[0] == 409
      2 assert res.values[1] == 156
      3 assert res.values[2] == 345

NameError: name 'res' is not defined
```